# `Coffee Problem "Fizzbuzz"`: `assignment subtitle`
_Version 0.0.1_

*All of the header information is important. Please read it..*

**Topics number of exercises:** This problem builds on your knowledge of `this, that, and the other`. It has **3** exercises numbered 0 to **2**. There are **total points** available points. However to earn 100% the threshold is **points cap** points. (Therefore once you hit **points cap** points you can stop. There is no extra credit for exceeding this threshold.)

**Exercise ordering:** Each exercise builds logically on previous exercises but you may solve them in any order. That is if you can't solve an exercise you can still move on and try the next one. Use this to your advantage as the exercises are **not** necessarily ordered in terms of difficulty. Higher point values generally indicate more difficult exercises. 

**Demo cells:** Code cells starting with the comment `### Run Me!!!` load results from prior exercises applied to the entire data set and use those to build demo inputs. These must be run for subsequent demos to work properly but they do not affect the test cells. The data loaded in these cells may be rather large (at least in terms of human readability). You are free to print or otherwise use Python to explore them but we may not print them in the starter code.

**Debugging your code:** Right before each exercise test cell there is a block of text explaining the variables available to you for debugging. You may use these to test your code and can print/display them as needed (careful when printing large objects you may want to print the head or chunks of rows at a time).

**Exercise point breakdown:**


- Exercise 0 - : **1** point(s)

- Exercise 1 - : **1** point(s)

- Exercise 2 - : **1** point(s)


**Final reminders:** 

- Submit after **every exercise**
- Review the generated grade report after you submit to see what errors were returned
- Stay calm, skip problems as needed and take short breaks at your leisure

```python
### Global imports
import dill
from cse6040_devkit import plugins, utils
import re 
import pandas as pd

utils.add_from_file('convert_to_str', plugins)
utils.add_from_file('foo', utils)

```

### Exercise 0: (1 points)
**fizzbuzz**  

**Your task:** define `fizzbuzz_soln` as follows:

`fizzbuzz_soln` solves the classic "fizzbuzz" problem for a single integer. 

**Inputs**:
- `i` (int): Any integer

**Returns**: one of the following depending on whether 3, 5, both or neither divide evenly into `i`
- `'Fizz'`     if i is divisible by 3                Example: 9
- `'Buzz'`     if i is divisible by 5                Example: 35
- `'FizzBuzz'` if i is divisible by both 3 and 5     Example: 45
- `i`          if i is divisible by neither 3 nor 5  Example: 22

---
We have defined a **helper function**, `x_divides_y` as follows:  

`x_divides_y` evaluates whether one integer (`x`) divides evenly into another integer (`y`).


```python
### Helper Function
def x_divides_y(x: int, y: int):
    return y%x == 0

### Solution - Exercise 0  
def fizzbuzz_soln(z: int):
    ### BEGIN SOLUTION
    return ("Fizz"*x_divides_y(3, z)+"Buzz"*x_divides_y(5, z)) or z
    ### END SOLUTION

### Demo function call
for x in [9, 35, 45, 22]:
    print(f'fizzbuzz_soln({x}) -> {fizzbuzz_soln(x)}')
```

 

**The demo should display this printed output.**
```
fizzbuzz_soln(9) -> Fizz
fizzbuzz_soln(35) -> Buzz
fizzbuzz_soln(45) -> FizzBuzz
fizzbuzz_soln(22) -> 22
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for fizzbuzz (exercise 0). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_0, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 0  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=plugins.convert_to_str(fizzbuzz_soln),
              ex_name='fizzbuzz',
              key=b'm_8VEe8rANbW9jOWIO0Rnu4v-gwlfczi_3yl_z0lBIo=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to fizzbuzz did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=plugins.convert_to_str(fizzbuzz_soln),
              ex_name='fizzbuzz',
              key=b'Mm4d7jm3-w_0wdLOp-1GLFB1IPLKtaq_apcVh4y-Xfg=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to fizzbuzz did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
free_data = utils.load_object_from_publicdata('free_data')

```

### Exercise 1: (1 points)
**freebie**  

**Example:** we have defined `freebie` as follows:



```python
### Solution - Exercise 1  
def freebie():
    ### This will never be tested
    

    print(free_data)
```

 
 <!-- Test Cell Boilerplate -->  
 The test cell below will always pass. Please submit to collect your free points for freebie (exercise 1).
 

<!-- grade_id: ex_1, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 1  


print('Passed! Please submit.')
```

```python
### Run Me!!!
char_df = utils.load_object_from_publicdata('char_df')
demo_result_species_count_TRUE = utils.load_object_from_publicdata('demo_result_species_count_TRUE')

```

### Exercise 2: (1 points)
**species_count**  

**Your task:** define `species_count_query` as follows:

The query should return these columns
- species (string)
- count (int): count of species value in characters

```python
### Solution - Exercise 2  

species_count_query = '''YOUR QUERY HERE'''
### BEGIN SOLUTION
species_count_query = '''
                      select species, count(1) as "count"
                      from characters
                      group by species
                  '''
### END SOLUTION

### Demo function call
from cse6040_devkit.tester_fw.test_utils \
    import dfs_to_conn
conn = dfs_to_conn({'characters': char_df})
demo_result_species_count = \
    plugins.sql_executor(species_count_query)(conn)
display(demo_result_species_count.head())
```

 

**The demo should display this output.**  

|    | species     |   count |
|---:|:------------|--------:|
|  0 | Ardennian   |       1 |
|  1 | Besalisk    |       1 |
|  2 | Chiss       |       2 |
|  3 | Clawdite    |       1 |
|  4 | Dathomirian |       3 |


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for species_count (exercise 2). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_2, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 2  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=plugins.sql_executor(species_count_query),
              ex_name='species_count',
              key=b'm_8VEe8rANbW9jOWIO0Rnu4v-gwlfczi_3yl_z0lBIo=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to species_count did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=plugins.sql_executor(species_count_query),
              ex_name='species_count',
              key=b'Mm4d7jm3-w_0wdLOp-1GLFB1IPLKtaq_apcVh4y-Xfg=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to species_count did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```