# `assignment name`: `assignment subtitle`
_Version 0.0.1_

*All of the header information is important. Please read it..*

**Topics number of exercises:** This problem builds on your knowledge of `this, that, and the other`. It has **11** exercises numbered 0 to **10**. There are **total points** available points. However to earn 100% the threshold is **points cap** points. (Therefore once you hit **points cap** points you can stop. There is no extra credit for exceeding this threshold.)

**Exercise ordering:** Each exercise builds logically on previous exercises but you may solve them in any order. That is if you can't solve an exercise you can still move on and try the next one. Use this to your advantage as the exercises are **not** necessarily ordered in terms of difficulty. Higher point values generally indicate more difficult exercises. 

**Demo cells:** Code cells starting with the comment `### Run Me!!!` load results from prior exercises applied to the entire data set and use those to build demo inputs. These must be run for subsequent demos to work properly but they do not affect the test cells. The data loaded in these cells may be rather large (at least in terms of human readability). You are free to print or otherwise use Python to explore them but we may not print them in the starter code.

**Debugging your code:** Right before each exercise test cell there is a block of text explaining the variables available to you for debugging. You may use these to test your code and can print/display them as needed (careful when printing large objects you may want to print the head or chunks of rows at a time).

**Exercise point breakdown:**


- Exercise 0 - : **1** point(s)

- Exercise 1 - : **1** point(s)

- Exercise 2 - : **1** point(s)

- Exercise 3 - : **1** point(s)

- Exercise 4 - : **1** point(s)

- Exercise 5 - : **1** point(s)

- Exercise 6 - : **1** point(s)

- Exercise 7 - : **1** point(s)

- Exercise 8 - : **1** point(s)

- Exercise 9 - : **1** point(s)

- Exercise 10 - : **1** point(s)


**Final reminders:** 

- Submit after **every exercise**
- Review the generated grade report after you submit to see what errors were returned
- Stay calm, skip problems as needed and take short breaks at your leisure

```python
### Global imports
import dill
from cse6040_devkit import plugins, utils
from collections import Counter
from pprint import pprint
from sqlite3 import connect
import pandas as pd 
import numpy as np 
from scipy.sparse import coo_matrix

utils.add_from_file('do_something', utils)

```

### Exercise 0: (1 points)
**inverted_dictionary**  

**Your task:** define `invert_dict` as follows:

Use input to create a new dictionary with keys and values swapped

There are configuration options to determine behavior when there is not a one-to-one mapping of keys to distinct, hashable values.

Args:
- d (dict): dictionary to invert
- allow_dup_vals (bool): 
    - When set to `True` the result will map any duplicated, hashable value in `d` to the `set` of its keys.
    - When set to `False` a `ValueError` is raised if there are any duplicate, hashable values in `d`.
- allow_unhashable_vals (bool): 
    -When set to `True` any unhashable values in `d` are not included in the result.  
    -When set to `False` a `ValueError` is raised if there are any unhashable values in `d`.

Raises:
- ValueError: if there are duplicate values in `d` and `allow_dup_vals` is `False`.
- ValueError: if there are unhashable values in `d` and `allow_unhashable_vals` is `False`.

Returns:
- dict: A new dictionary mapping the values of `d` to their associated keys.  


```python
### Solution - Exercise 0  
def invert_dict(d: dict, allow_dup_vals: bool, allow_unhashable_vals: bool) -> dict:
    ### BEGIN SOLUTION
    d_out = {}
    for v, k in d.items():
        if not utils.is_hashable(k):
            if allow_unhashable_vals:
                continue
            else:
                raise ValueError
        if k in d_out:
            if allow_dup_vals:
                if isinstance(d_out[k], set):
                    d_out[k].add(v)
                else:
                    d_out[k] = {d_out[k], v}
                pass
            else:
                raise ValueError
        else:
            d_out[k] = v
    return d_out
    ### END SOLUTION

### Demo function call
# We use an error handler to catch errors raised by your solution
# 
wrapped = plugins.error_handler(invert_dict)
hashable_no_dups = {1:2, 3:4, 5:6}
unhashable_no_dups = {1:[3,4,5], 2:3, 4:5, 7:[3,4,5]}
unhashable_with_dups = {1:[3,4,5], 2:3, 4:5}
hashable_with_dups = {1:2, 3:4, 5:2}
test_grid = [[True, True], [True, False], [False, True], [False, False]]

for tgi in test_grid:
    print('tgi = ', dict(zip(['allow_dup_vals', 'allow_unhashable_vals'], tgi)))
    print(f"hashable_no_dups result: {wrapped(hashable_no_dups, *tgi)}")
    print(f"unhashable_no_dups result: {wrapped(unhashable_no_dups, *tgi)}")
    print(f"unhashable_with_dups result: {wrapped(unhashable_with_dups, *tgi)}")
    print(f"hashable_with_dups result: {wrapped(hashable_with_dups, *tgi)}")
    print()
```

 There are 4 possible combinations of having duplicates and being hashable. This means that there are possible states for the data and four possible states for the parameters. We have created a test grid with simple examples of all combinations.  


**The demo should display this printed output.**
```
tgi =  {'allow_dup_vals': True, 'allow_unhashable_vals': True}
hashable_no_dups result: (False, {2: 1, 4: 3, 6: 5})
unhashable_no_dups result: (False, {3: 2, 5: 4})
unhashable_with_dups result: (False, {3: 2, 5: 4})
hashable_with_dups result: (False, {2: {1, 5}, 4: 3})

tgi =  {'allow_dup_vals': True, 'allow_unhashable_vals': False}
hashable_no_dups result: (False, {2: 1, 4: 3, 6: 5})
unhashable_no_dups result: (True, None)
unhashable_with_dups result: (True, None)
hashable_with_dups result: (False, {2: {1, 5}, 4: 3})

tgi =  {'allow_dup_vals': False, 'allow_unhashable_vals': True}
hashable_no_dups result: (False, {2: 1, 4: 3, 6: 5})
unhashable_no_dups result: (False, {3: 2, 5: 4})
unhashable_with_dups result: (False, {3: 2, 5: 4})
hashable_with_dups result: (True, None)

tgi =  {'allow_dup_vals': False, 'allow_unhashable_vals': False}
hashable_no_dups result: (False, {2: 1, 4: 3, 6: 5})
unhashable_no_dups result: (True, None)
unhashable_with_dups result: (True, None)
hashable_with_dups result: (True, None)
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for inverted_dictionary (exercise 0). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_0, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 0  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=plugins.error_handler(invert_dict),
              ex_name='inverted_dictionary',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to inverted_dictionary did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=plugins.error_handler(invert_dict),
              ex_name='inverted_dictionary',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to inverted_dictionary did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
demo_result_concatenate_quotes_TRUE = utils.load_object_from_publicdata('demo_result_concatenate_quotes_TRUE')

```

### Exercise 1: (1 points)
**concatenate_quotes**  

**Your task:** define `concatenate_quotes_query` as follows:

Re-organize the data in the `quotes` table such that all `quote` values associated with a particular `character_name` are in a single row.  

Result columns:  
- `name` - The name of a character
- `quotes` - All quote values associated with the character, concatenated together.
    - Quotes should be ordered alphabetically.
    - Quotes should be separated by the "pipe" character `'|'`.
- `quote_count` - The number of quotes associated with the character.

**The result should be sorted from highest-to-lowest with respect to `quote_count` with ties being sorted alphabetically by `name`.**

The quotes table has at least the following columns indicating quotations spoken by a character:

- character_name (varchar): Name of the character
- quote (varchar): A film quotation



```python
### Solution - Exercise 1  

concatenate_quotes_query = '''YOUR QUERY HERE'''
### BEGIN SOLUTION
concatenate_quotes_query = '''
                        SELECT 
                            character_name AS name,
                            GROUP_CONCAT(quote, '|') AS quotes,
                            COUNT(*) AS quote_count
                        FROM (SELECT * FROM quotes order by quote)
                        GROUP BY character_name
                        ORDER BY quote_count DESC, name'''
### END SOLUTION

### Demo function call
with connect('resource/asnlib/publicdata/star_wars.db') as demo_conn:
    demo_result_concatenate_quotes = plugins.sql_executor(concatenate_quotes_query)(demo_conn)
utils.display_df_text_wrap(demo_result_concatenate_quotes)
```

 

**The demo should display this output.**  

|    | name              | quotes                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                              |   quote_count |
|---:|:------------------|:--------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------|--------------:|
|  0 | Han Solo          | Boba Fett? Boba Fett? Where?\|Chewie, we’re home.\|Good against remotes is one thing. Good against the living, that’s something else.\|Great, kid. Don’t get cocky.\|I know.\|I think I just blasted it.\|I think you just can’t bear to let a gorgeous guy like me out of your sight.\|I thought they smelled bad on the outside.\|I've got a bad feeling about this.\|It’s not my fault.\|I’ll see you in hell!\|I’m out of it for a little while, everyone gets delusions of grandeur.\|I’ve got a bad feeling about this.\|Laugh it up, fuzzball.\|Never tell me the odds!\|Who’s scruffy-looking?\|You’re all clear, kid. Now let’s blow this thing and go home.\|You’re all clear, kid. Now let’s blow this thing and go home.                                                                                                                                |            18 |
|  1 | Obi-Wan Kenobi    | I have a bad feeling about this.\|I will do what I must.\|In my experience, there is no such thing as luck.\|In my experience, there's no such thing as luck.\|Remember, the Force will be with you, always.\|Strike me down, and I will become more powerful than you could possibly imagine.\|That’s no moon. It’s a space station.\|The Force will be with you, always.\|The Force will be with you, always.\|The Force will be with you. Always.\|These aren’t the droids you’re looking for.\|This is the weapon of a Jedi Knight. Not as clumsy or random as a blaster; an elegant weapon for a more civilized age.\|You were right about one thing, Master. The negotiations were short.\|You were the Chosen One!\|You will find that many of the truths we cling to depend greatly on our own point of view.\|Your eyes can deceive you; don’t trust them. |            16 |
|  2 | Darth Vader       | I am altering the deal. Pray I don’t alter it any further.\|I am your father.\|I am your father.\|I am your father.\|I find your lack of faith disturbing.\|I find your lack of faith disturbing.\|I find your lack of faith disturbing.\|I’ve been waiting for you, Obi-Wan. We meet again, at last.\|Just for once, let me look on you with my own eyes.\|The Force is strong with this one.\|The circle is now complete.\|When I left you, I was but the learner. Now I am the master.\|You don’t know the power of the dark side.\|You have failed me for the last time, Admiral.                                                                                                                                                                                                                                                                               |            14 |
|  3 | Luke Skywalker    | Do you know what I’ve become?\|I am a Jedi, like my father before me.\|I am a Jedi, like my father before me.\|I am a Jedi, like my father before me.\|I’ll never turn to the dark side.\|I’ll never turn to the dark side.\|I’ve got a bad feeling about this.\|May the Force be with you.\|My name is Luke Skywalker. I’m here to rescue you.\|You have failed, your highness. I am a Jedi, like my father before me.\|You serve your master well. And you will be rewarded.\|Your overconfidence is your weakness.                                                                                                                                                                                                                                                                                                                                               |            12 |
|  4 | Yoda              | A Jedi uses the Force for knowledge and defense, never for attack.\|Adventure. Excitement. A Jedi craves not these things.\|Do or do not, there is no try.\|Fear is the path to the dark side.\|Fear leads to anger, anger leads to hate, hate leads to suffering.\|Judge me by my size, do you?\|Luminous beings are we, not this crude matter.\|Once you start down the dark path, forever will it dominate your destiny.\|There is no emotion, there is peace.\|Wars not make one great.\|You must unlearn what you have learned.                                                                                                                                                                                                                                                                                                                                |            11 |
|  5 | Leia Organa       | Help me, Obi-Wan Kenobi. You're my only hope.\|I can’t believe he’s gone.\|I love you.\|Why, you stuck-up, half-witted, scruffy-looking Nerf herder!\|Would somebody get this big walking carpet out of my way?\|You do have your moments. Not many, but you have them.\|You do have your moments. Not many, but you have them.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                     |             7 |
|  6 | Emperor Palpatine | Now, young Skywalker, you will die.\|Power! Unlimited power!\|So be it, Jedi.\|The dark side of the Force is a pathway to many abilities some consider to be unnatural.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                             |             4 |
|  7 | Anakin Skywalker  | I can’t watch anymore.\|This is where the fun begins.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                               |             2 |
|  8 | C-3PO             | We seem to be made to suffer. It’s our lot in life.\|We’re doomed.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                  |             2 |
|  9 | Lando Calrissian  | I’ve just made a deal that will keep the Empire out of here forever.\|You look absolutely beautiful. You truly belong here with us among the clouds.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                |             2 |
| 10 | Padmé Amidala     | I will not let this Republic that has stood for a thousand years be split in two.\|So this is how liberty dies. With thunderous applause.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           |             2 |
| 11 | Qui-Gon Jinn      | There’s always a bigger fish.\|Your focus determines your reality.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                  |             2 |
| 12 | Admiral Ackbar    | It's a trap!                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                        |             1 |
| 13 | Admiral Piett     | Bounty hunters. We don’t need their scum.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           |             1 |
| 14 | Boba Fett         | He’s no good to me dead.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                            |             1 |
| 15 | Count Dooku       | I sense great fear in you, Skywalker.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                               |             1 |
| 16 | Gold Five         | Stay on target.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                     |             1 |
| 17 | Opening Crawl     | A long time ago in a galaxy far, far away…                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                          |             1 |
| 18 | Shmi Skywalker    | You can’t stop the change, any more than you can stop the suns from setting.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                        |             1 |


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for concatenate_quotes (exercise 1). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_1, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 1  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=plugins.sql_executor(concatenate_quotes_query),
              ex_name='concatenate_quotes',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to concatenate_quotes did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=plugins.sql_executor(concatenate_quotes_query),
              ex_name='concatenate_quotes',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to concatenate_quotes did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
demo_result_extract_article_words_TRUE = utils.load_object_from_publicdata('demo_result_extract_article_words_TRUE')

```

### Exercise 2: (1 points)
**extract_article_words**  

**Your task:** define `extract_article_words` as follows:



Given a dictionary, `record`, with the key `'Article'` mapped to a string of unstructured text, create a Python set containing all **distinct**, **lower-case** "words" in the text. For our purposes a "word" is _a sequence of **one or more** consecutive **letters**._

```python
### Solution - Exercise 2  
def extract_article_words(record: dict) -> set:
    ### BEGIN SOLUTION
    import re
    return set(re.sub('[^\w\s]|[\d_]', ' ', record['Article'].lower()).split())
    ### END SOLUTION

### Demo function call
demo_record = {'Article': 'This is a "news article". \n\tNum-bers, punctuati0n, and under_scores aren\'t letters, so they cannot be part of a word!',
    'Date': '2/9/2015',
    'Heading': 'hsbc admits swiss bank failings over client tax',
    'NewsType': 'business'}
demo_result_extract_article_words = extract_article_words(demo_record)
print(demo_result_extract_article_words)

```

 

**The demo should display this printed output.**
```
{'punctuati', 'word', 'this', 'n', 'of', 'article', 'and', 't', 'part', 'they', 'cannot', 'under', 'news', 'aren', 'be', 'so', 'is', 'num', 'bers', 'scores', 'a', 'letters'}
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for extract_article_words (exercise 2). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_2, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 2  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=extract_article_words,
              ex_name='extract_article_words',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to extract_article_words did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=extract_article_words,
              ex_name='extract_article_words',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to extract_article_words did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
demo_result_count_articles_TRUE = utils.load_object_from_publicdata('demo_result_count_articles_TRUE')

```

### Exercise 3: (1 points)
**count_articles**  

**Your task:** define `count_articles` as follows:

Update a running count by one for each distinct word in an article.  

Args:  
- running_count (Counter): Maps words to the number of articles they have appeared in.  
- article_words (set): The distinct words appearing in a single article.  

Returns:  
- Counter: Updated running total.


```python
### Solution - Exercise 3  
def count_articles(running_count: Counter, article_words):
    ### BEGIN SOLUTION
    return running_count + Counter(article_words)
    ### END SOLUTION

### Demo function call
demo_running_count = Counter({'georgia': 3, 'tech': 5, 'jacket': 15})
demo_article_words = {'ga', 'tech', 'yellow', 'jacket'}
demo_result_count_articles = count_articles(demo_running_count, demo_article_words)
pprint(demo_result_count_articles)
```

 

**The demo should display this output.**  

Counter({'jacket': 16, 'tech': 6, 'georgia': 3, 'yellow': 1, 'ga': 1})


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for count_articles (exercise 3). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_3, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 3  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=count_articles,
              ex_name='count_articles',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to count_articles did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=count_articles,
              ex_name='count_articles',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to count_articles did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

### Map and Reduce

The widely used Hadoop and Spark ecosystems use a concept called map-reduce to parallelize big data operations. The idea is to apply a transformation to each item in a data column then combine each transformed item into a running total value. This allows breaking up data across multiple machines for processing (map), computing a running total value for each machine (reduce), and finally combining the running total values from all the machines into a single result.   

Consider this simple mapping function. It takes a single string argument and returns the same string in lower-case:

```
def map_fun(item):
    assert isinstance(item, str), 'All items must be strings'
    return item.lower()
```

If we apply `map_fun` to the list `demo_some_iter = ['This', 'is', 'An', 'ExAmpL3']`, we expect to get `['this', 'is', 'an', 'exampl3']`.

Now consider this reducing function. It takes two arguments, `value` - which is the "running total value" and `item` which is the transformed value being combined into `value`. It returns the new "running total value".  

```
def red_fun(value, item):
    assert isinstance(value, str), 'value must be a string'
    return value + '|' + item
```

If we iterate over the items in our intermediate result updating our running total value using something like `value = red_fun(value, item)` in each iteration, we would get these snapshot results on each iteration:

- `'this'`
- `'this|is'`
- `'this|is|an'`
- `'this|is|an|exampl3'`

```python
### Run Me!!!
demo_result_record_examiner_TRUE = utils.load_object_from_publicdata('demo_result_record_examiner_TRUE')

```

### Exercise 4: (1 points)
**record_examiner**  

**Your task:** define `record_examiner` as follows:

Applies the `map_func` to each item in `some_iter` combines the results with `red_func`.

Args:
- some_iter (iterable): An iterable object.  
- map_func (function): A function that takes a single input and returns a single output.  
- red_func (function): A function which takes two inputs (`value` and `item`). It returns the result of "combining" `item` with `value`.
- initial (any, optional): The initial `value` used in `red_func`. Defaults to None, which will not use an initial value.

Returns:
    any: The result of applying `map_func` to each item in `some_iter` and then combining them with `red_func`.


```python
### Solution - Exercise 4  
def record_examiner(some_iter, map_func, red_func, initial=None):
    ### BEGIN SOLUTION
    from functools import reduce
    if initial is not None:
        return reduce(red_func, map(map_func, some_iter), initial)
    else:
        return reduce(red_func, map(map_func, some_iter))
    ### END SOLUTION

### Demo function call
def insert_in_sorted_list(L, item):
    '''Inserts an item into L such that L remains sorted
    '''
    import bisect
    bisect.insort(L, item)
    return L

demo_input_vars = {
    'some_iter': ['This', 'is', 'An', 'ExAmpL3'],
    'map_func': lambda s: s.lower(),
    'red_func': insert_in_sorted_list,
    'initial': [] 
}
demo_original_input_vars = {
    'some_iter': ['This', 'is', 'An', 'ExAmpL3']        
}
demo_result_record_examiner = record_examiner(**demo_input_vars)
print(demo_result_record_examiner)

```

 

**The demo should display this printed output.**
```
['an', 'exampl3', 'is', 'this']
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for record_examiner (exercise 4). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_4, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 4  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=record_examiner,
              ex_name='record_examiner',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to record_examiner did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=record_examiner,
              ex_name='record_examiner',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to record_examiner did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
demo_result_conn_to_df_TRUE = utils.load_object_from_publicdata('demo_result_conn_to_df_TRUE')

```

### Exercise 5: (1 points)
**conn_to_df**  

**Your task:** define `concatenate_quotes_conn_to_df` as follows:

Re-organize the data in the `quotes` table such that all `quote` values associated with a particular `character_name` are in a single row.  

Args:

- conn (sqlite3 connection): Connection with table called `quotes`.
    - `quotes` will have `character_name` and `quote` columns attributing a film quote to a character. Each record is one film quote.

Returns:
- DataFrame: contains the following columns:
    - `name` - The name of a character
    - `quotes` - All quote values associated with the character, concatenated together.
        - Quotes should be ordered alphabetically.
        - Quotes should be separated by the "pipe" character `'|'`.
    - `quote_count` - The number of quotes associated with the character. 


```python
### Solution - Exercise 5  
def concatenate_quotes_conn_to_df(conn):
    ### BEGIN SOLUTION
    df = pd.read_sql('''
                        SELECT 
                            character_name AS name,
                            GROUP_CONCAT(quote, '|') AS quotes,
                            COUNT(*) AS quote_count
                        FROM (SELECT * FROM quotes order by quote)
                        GROUP BY character_name
                        ORDER BY quote_count DESC, name''',
                    conn)
    df = df.astype({'name': str,
               'quotes': str,
               'quote_count': int})
    return df
    ### END SOLUTION

### Demo function call
with connect('resource/asnlib/publicdata/star_wars.db') as demo_conn:
    demo_result_conn_to_df = concatenate_quotes_conn_to_df(demo_conn)
utils.display_df_text_wrap(demo_result_conn_to_df)
```

 

**The demo should display this output.**  

|    | name              | quotes                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                              |   quote_count |
|---:|:------------------|:--------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------|--------------:|
|  0 | Han Solo          | Boba Fett? Boba Fett? Where?\|Chewie, we’re home.\|Good against remotes is one thing. Good against the living, that’s something else.\|Great, kid. Don’t get cocky.\|I know.\|I think I just blasted it.\|I think you just can’t bear to let a gorgeous guy like me out of your sight.\|I thought they smelled bad on the outside.\|I've got a bad feeling about this.\|It’s not my fault.\|I’ll see you in hell!\|I’m out of it for a little while, everyone gets delusions of grandeur.\|I’ve got a bad feeling about this.\|Laugh it up, fuzzball.\|Never tell me the odds!\|Who’s scruffy-looking?\|You’re all clear, kid. Now let’s blow this thing and go home.\|You’re all clear, kid. Now let’s blow this thing and go home.                                                                                                                                |            18 |
|  1 | Obi-Wan Kenobi    | I have a bad feeling about this.\|I will do what I must.\|In my experience, there is no such thing as luck.\|In my experience, there's no such thing as luck.\|Remember, the Force will be with you, always.\|Strike me down, and I will become more powerful than you could possibly imagine.\|That’s no moon. It’s a space station.\|The Force will be with you, always.\|The Force will be with you, always.\|The Force will be with you. Always.\|These aren’t the droids you’re looking for.\|This is the weapon of a Jedi Knight. Not as clumsy or random as a blaster; an elegant weapon for a more civilized age.\|You were right about one thing, Master. The negotiations were short.\|You were the Chosen One!\|You will find that many of the truths we cling to depend greatly on our own point of view.\|Your eyes can deceive you; don’t trust them. |            16 |
|  2 | Darth Vader       | I am altering the deal. Pray I don’t alter it any further.\|I am your father.\|I am your father.\|I am your father.\|I find your lack of faith disturbing.\|I find your lack of faith disturbing.\|I find your lack of faith disturbing.\|I’ve been waiting for you, Obi-Wan. We meet again, at last.\|Just for once, let me look on you with my own eyes.\|The Force is strong with this one.\|The circle is now complete.\|When I left you, I was but the learner. Now I am the master.\|You don’t know the power of the dark side.\|You have failed me for the last time, Admiral.                                                                                                                                                                                                                                                                               |            14 |
|  3 | Luke Skywalker    | Do you know what I’ve become?\|I am a Jedi, like my father before me.\|I am a Jedi, like my father before me.\|I am a Jedi, like my father before me.\|I’ll never turn to the dark side.\|I’ll never turn to the dark side.\|I’ve got a bad feeling about this.\|May the Force be with you.\|My name is Luke Skywalker. I’m here to rescue you.\|You have failed, your highness. I am a Jedi, like my father before me.\|You serve your master well. And you will be rewarded.\|Your overconfidence is your weakness.                                                                                                                                                                                                                                                                                                                                               |            12 |
|  4 | Yoda              | A Jedi uses the Force for knowledge and defense, never for attack.\|Adventure. Excitement. A Jedi craves not these things.\|Do or do not, there is no try.\|Fear is the path to the dark side.\|Fear leads to anger, anger leads to hate, hate leads to suffering.\|Judge me by my size, do you?\|Luminous beings are we, not this crude matter.\|Once you start down the dark path, forever will it dominate your destiny.\|There is no emotion, there is peace.\|Wars not make one great.\|You must unlearn what you have learned.                                                                                                                                                                                                                                                                                                                                |            11 |
|  5 | Leia Organa       | Help me, Obi-Wan Kenobi. You're my only hope.\|I can’t believe he’s gone.\|I love you.\|Why, you stuck-up, half-witted, scruffy-looking Nerf herder!\|Would somebody get this big walking carpet out of my way?\|You do have your moments. Not many, but you have them.\|You do have your moments. Not many, but you have them.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                     |             7 |
|  6 | Emperor Palpatine | Now, young Skywalker, you will die.\|Power! Unlimited power!\|So be it, Jedi.\|The dark side of the Force is a pathway to many abilities some consider to be unnatural.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                             |             4 |
|  7 | Anakin Skywalker  | I can’t watch anymore.\|This is where the fun begins.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                               |             2 |
|  8 | C-3PO             | We seem to be made to suffer. It’s our lot in life.\|We’re doomed.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                  |             2 |
|  9 | Lando Calrissian  | I’ve just made a deal that will keep the Empire out of here forever.\|You look absolutely beautiful. You truly belong here with us among the clouds.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                |             2 |
| 10 | Padmé Amidala     | I will not let this Republic that has stood for a thousand years be split in two.\|So this is how liberty dies. With thunderous applause.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           |             2 |
| 11 | Qui-Gon Jinn      | There’s always a bigger fish.\|Your focus determines your reality.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                  |             2 |
| 12 | Admiral Ackbar    | It's a trap!                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                        |             1 |
| 13 | Admiral Piett     | Bounty hunters. We don’t need their scum.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           |             1 |
| 14 | Boba Fett         | He’s no good to me dead.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                            |             1 |
| 15 | Count Dooku       | I sense great fear in you, Skywalker.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                               |             1 |
| 16 | Gold Five         | Stay on target.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                     |             1 |
| 17 | Opening Crawl     | A long time ago in a galaxy far, far away…                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                          |             1 |
| 18 | Shmi Skywalker    | You can’t stop the change, any more than you can stop the suns from setting.                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                        |             1 |


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for conn_to_df (exercise 5). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_5, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 5  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=concatenate_quotes_conn_to_df,
              ex_name='conn_to_df',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to conn_to_df did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=concatenate_quotes_conn_to_df,
              ex_name='conn_to_df',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to conn_to_df did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
char_df = utils.load_object_from_publicdata('char_df')
spec_df = utils.load_object_from_publicdata('spec_df')
demo_result_species_count_TRUE = utils.load_object_from_publicdata('demo_result_species_count_TRUE')

```

### Exercise 6: (1 points)
**species_count**  

**Your task:** define `species_count` as follows:

Determine the count of characters in `char_df` belonging to each `classification` in `spec_df`.

Args:
- `char_df` (pd.DataFrame): contains character data including the 'name' and 'species' columns.
- `spec_df` (pd.DataFrame): contains species data including the 'name' and 'classification' columns. The values in `spec_df['name']` refer to the same data as `char_df['species']`.

Returns:
- pd.DataFrame: contains the following columns:
    - `classification` (str): species classification. Each row should have a distinct value in this column.
    - `count` (int): count indicating how many characters from the `char_df` are members of a species with the `classification`.


```python
### Solution - Exercise 6  
def species_count(char_df, spec_df):
    ### BEGIN SOLUTION
    ### BEGIN SOLUTION
    merged = char_df.merge(spec_df, left_on='species', right_on='name')
    return merged['classification'].value_counts().reset_index().set_axis(['classification', 'count'], axis=1)
    ### END SOLUTION
    ### END SOLUTION

### Demo function call
demo_result_species_count = species_count(char_df, spec_df)
utils.display_df_text_wrap(demo_result_species_count)
```

 

**The demo should display this output.**  

|    | classification   |   count |
|---:|:-----------------|--------:|
|  0 | Mammal           |      74 |
|  1 | Amphibian        |       4 |
|  2 | Artificial       |       4 |
|  3 | Reptilian        |       3 |
|  4 | Hybrid           |       3 |
|  5 | Unknown          |       1 |
|  6 | Gastropod        |       1 |


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for species_count (exercise 6). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_6, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 6  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=species_count,
              ex_name='species_count',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to species_count did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=species_count,
              ex_name='species_count',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to species_count did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

### Exercise 7: (1 points)
**count_values**  

**Your task:** define `count_values` as follows:

Get the count of each distinct value in a Pandas Series  

Args:  

- s (pd.Series): Any Pandas Series

Return:  

- pd.Series: The index of the result series is all distinct values from the input Series, `s`. The values of the result series are the occurrance counts of the index within the input Series, `s`.


```python
### Run Me!!!
demo_result_count_values_TRUE = utils.load_object_from_publicdata('demo_result_count_values_TRUE')

```

```python
### Solution - Exercise 7  
def count_values(s):
    ### BEGIN SOLUTION
    return s.value_counts()
    ### END SOLUTION

### Demo function call
demo_s = pd.Series([1,2,3,3,2,2,2,1,1,1,2,3,1,2,3,1,2,3])
demo_result_count_values = count_values(demo_s)
print(demo_result_count_values)

```

 

**The demo should display this printed output.**
```
2    7
1    6
3    5
Name: count, dtype: int64
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for count_values (exercise 7). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_7, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 7  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=count_values,
              ex_name='count_values',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to count_values did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=count_values,
              ex_name='count_values',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to count_values did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
demo_result_mat_vec_div_TRUE = utils.load_object_from_publicdata('demo_result_mat_vec_div_TRUE')

```

### Exercise 8: (1 points)
**mat_vec_div**  

**Your task:** define `mat_vec_div` as follows:

Compute matrix-vector product of $Ax^{-1}$  

Args:

- A (np.ndarray): A 2-D numeric array
- x (np.ndarray): A 1-D numeric array

Returns:

- np.ndarray: The matrix-vector product of `A` and _the inverse of_ `x`.


Mathematically, the inverse of a vector,  $x$, is the vector, $x^{-1}$, such that $x^{-1}x = 1$. You should be able to convince yourself that  

$$x^{-1} = \frac{x}{\|x\|^2}$$

For example, consider when $x$ is $[1, 2]$. This makes $\|x\|^2 = 1^2 + 2^2 = 5$. This makes $x^{-1} = [0.2, 0.4]$

```python
### Solution - Exercise 8  
def mat_vec_div(A, x):
    ### BEGIN SOLUTION
    x_inv = x / x.dot(x.T)
    return A.dot(x_inv).round(3)
    ### END SOLUTION

### Demo function call
demo_A = np.array([[1,2], [7,4]])
demo_x = np.array([1,2])

demo_result_mat_vec_div = mat_vec_div(demo_A, demo_x)
print(demo_result_mat_vec_div)

```

 

**The demo should display this printed output.**
```
[1. 3.]
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for mat_vec_div (exercise 8). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_8, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 8  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=mat_vec_div,
              ex_name='mat_vec_div',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to mat_vec_div did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=mat_vec_div,
              ex_name='mat_vec_div',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to mat_vec_div did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

```python
### Run Me!!!
demo_result_df_to_coo_TRUE = utils.load_object_from_publicdata('demo_result_df_to_coo_TRUE')

```

### Exercise 9: (1 points)
**df_to_coo**  

**Your task:** define `df_to_coo` as follows:


Create a COO matrix based off of a DataFrame

Args:
- df (DataFrame): A DataFrame representing a data matrix.
    - `data` is the value at position (`row`, `col) in the matrix.
- shape (tuple): Shape of the data matrix

Returns:
- scipy.sparse.coo_matrix: The matrix given by `df`.


```python
### Solution - Exercise 9  
def df_to_coo(df, shape):
    ### BEGIN SOLUTION
    return coo_matrix((df['data'], (df['row'], df['col'])), shape)
    ### END SOLUTION

### Demo function call
demo_shape = (33, 45)
demo_df = pd.DataFrame({'data': [2,3,4],
                        'row': [2, 20, 10],
                        'col': [4, 1, 30]})
demo_result_df_to_coo = df_to_coo(demo_df, demo_shape)
print(demo_result_df_to_coo)
print(f'{demo_result_df_to_coo.shape=}')
print(f'{demo_result_df_to_coo.data=}')
print(f'{demo_result_df_to_coo.row=}')
print(f'{demo_result_df_to_coo.col=}')

```

 

**The demo should display this printed output.**
```
(2, 4)	2
  (20, 1)	3
  (10, 30)	4
demo_result_df_to_coo.shape=(33, 45)
demo_result_df_to_coo.data=array([2, 3, 4])
demo_result_df_to_coo.row=array([ 2, 20, 10], dtype=int32)
demo_result_df_to_coo.col=array([ 4,  1, 30], dtype=int32)
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for df_to_coo (exercise 9). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_9, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 9  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=plugins.coo_plugin(df_to_coo),
              ex_name='df_to_coo',
              key=b'6AfuzbXbZoimCVSPv56sgGDDXYXLpbmJ70319bJ5i54=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to df_to_coo did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=plugins.coo_plugin(df_to_coo),
              ex_name='df_to_coo',
              key=b'b3LurxufK6uXEsHMqKTylu4CvoYNzO0Oa6ByQxb5P3s=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to df_to_coo did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

### Exercise 10: (1 points)
**load_the_data**  

**Example:** we have defined `load_the_data` as follows:



```python
### Helper Function
def do_a_thing():
    print('I did a thing!')

### Solution - Exercise 10  
def load_the_data(data):
    do_a_thing()
    utils.do_something()
    print(f'''{str(data)} loaded!''')
```

 
 <!-- Test Cell Boilerplate -->  
 The test cell below will always pass. Please submit to collect your free points for load_the_data (exercise 10).
 

<!-- grade_id: ex_10, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 10  


print('Passed! Please submit.')
```