```python
### Global imports
import dill
from cse6040_devkit import plugins, utils
import re 
import pandas as pd

utils.add_from_file('print_handler', plugins)

```

Output:
```
cse6040_devkit.plugins
```

```python
### Run Me!!!
demo_result_gt_comparison_TRUE = utils.load_object_from_publicdata('demo_result_gt_comparison_TRUE')

```

Output:
```
Successfully loaded demo_result_gt_comparison_TRUE.
```

### Exercise 0: (1 points)
**gt_comparison**  

**Your task:** define `x_gt_y` as follows:

Return `True` if `x` is greater than `y`. `False` otherwise.


```python
### Solution - Exercise 0  
def x_gt_y(x, y):
    ### BEGIN SOLUTION
    return x > y
    ### END SOLUTION

### Demo function call
x = 3
demo_result_x_gt_y = []
for y in range(2,6):
    result = x_gt_y(x,y)
    print(f'{x=}; {y=}; {result=}')
    demo_result_x_gt_y.append({
        'x': x,
        'y': y, 
        'result': result})

```

Output:
```
x=3; y=2; result=True
x=3; y=3; result=False
x=3; y=4; result=False
x=3; y=5; result=False
```

 

**The demo should display this printed output.**
```
x=3; y=2; result=True
x=3; y=3; result=False
x=3; y=4; result=False
x=3; y=5; result=False
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for gt_comparison (exercise 0). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_0, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 0  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=x_gt_y,
              ex_name='gt_comparison',
              key=b'nhnrMaKO7HzEE2viKVn7IQEKJwExvuVc3Jv_vK3ESAY=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to gt_comparison did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=x_gt_y,
              ex_name='gt_comparison',
              key=b'nHLHbiY7qwciZVnOK0Fuq2EkbQYIqU7DZ67sbGKHrQI=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to gt_comparison did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

Output:
```
gt_comparison test ran 99 iterations in 0.03 seconds
gt_comparison test ran 0 iterations in 0.01 seconds
Passed! Please submit.
```

 

**The demo should display this printed output.**
```
f(1) = 3(1) + 3 = 6

f(8) = 4(8) + -0.5 = 31.5

f(12) = 9(12) + 140 = 248
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for lin_printer_fac (exercise 1). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


```python
### Run Me!!!
demo_result_lin_print_fac_TRUE = utils.load_object_from_publicdata('demo_result_lin_print_fac_TRUE')

```

Output:
```
Successfully loaded demo_result_lin_print_fac_TRUE.
```

### Exercise 1: (10 points)
**lin_print_fac**  

**Your task:** define `lin_printer_fac` as follows:



```python
### Solution - Exercise 1  
def lin_printer_fac(m, b):
    ### BEGIN SOLUTION
    def f(x):
        y = m*x + b
        print(f'f({x}) = {m}({x}) + {b} = {y}')
        return y
    return f
    ### END SOLUTION

### Demo function call
wp_lin_printer_fac = plugins.print_handler(lin_printer_fac,
                                           show_printed_output=True)
### three examples
demo_result_lin_printer_fac = [
    wp_lin_printer_fac(m=3, b=3, x=1),
    wp_lin_printer_fac(m=4, b=-0.5, x=8),
    wp_lin_printer_fac(m=9, b=140, x=12)
]

```

Output:
```
f(1) = 3(1) + 3 = 6

f(8) = 4(8) + -0.5 = 31.5

f(12) = 9(12) + 140 = 248
```

 

**The demo should display this printed output.**
```
f(1) = 3(1) + 3 = 6

f(8) = 4(8) + -0.5 = 31.5

f(12) = 9(12) + 140 = 248
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for lin_print_fac (exercise 1). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_1, points: 10, grade: true, solution: false -->
```python
### Test Cell - Exercise 1  

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

    plugin_kwargs = utils.load_object_from_publicdata('lin_print_fac_plugin_kwargs')

# Execute test
passed, test_case_vars, e = execute_tests(func=plugins.print_handler(lin_printer_fac, **plugin_kwargs),
              ex_name='lin_print_fac',
              key=b'nhnrMaKO7HzEE2viKVn7IQEKJwExvuVc3Jv_vK3ESAY=', 
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to lin_print_fac did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=plugins.print_handler(lin_printer_fac, **plugin_kwargs),
              ex_name='lin_print_fac',
              key=b'nHLHbiY7qwciZVnOK0Fuq2EkbQYIqU7DZ67sbGKHrQI=', 
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to lin_print_fac did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

Output:
```
Successfully loaded lin_print_fac_plugin_kwargs.
lin_print_fac test ran 99 iterations in 0.01 seconds
lin_print_fac test ran 0 iterations in 0.01 seconds
Passed! Please submit.
```